How bats exit a crowded colony when relying on echolocation only - a modeling approach

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

Mazar & Yovel 2025 dissect the inverse problem of how echolocators in groups manage to navigate their surroundings despite intense jamming using computational simulations.

The authors show that despite the 'noisy' sensory environments that echolocating groups present, agents can still access some amount of echo-related information and use it to navigate their local environment. It is known that echolocating bats have strong small and large-scale spatial memory that plays an important role for individuals. The results from this paper also point to the potential importance of an even lower-level, short-term role of memory in the form of echo 'integration' across multiple calls, despite the unpredictability of echo detection in groups. The paper generates a useful basis to think about the mechanisms in echolocating groups for experimental investigations too.

Strengths:

The paper builds on biologically well-motivated and parametrised 2D acoustics and sensory simulation setup to investigate the various key parameters of interest

The 'null-model' of echolocators not being able to tell apart objects & conspecifics while echolocating still shows agents succesfully emerge from groups - even though the probability of emergence drops severely in comparison to cognitively more 'capable' agents. This is nonetheless an important result showing the direction-of-arrival of a sound itself is the 'minimum' set of ingredients needed for echolocators navigating their environment.

The results generate an important basis in unraveling how agents may navigate in sensorially noisy environments with a lot of irrelevant and very few relevant cues.

The 2D simulation framework is simple and computationally tractable enough to perform multiple runs to investigate many variables - while also remaining true to the aim of the investigation.

Weaknesses:

Authors have not yet provided convincing justification for the use of different echolocation phases during emergence and in cave behaviour. In the previous modelling paper cited for the details - here the bat-agents are performing a foraging task, and so the switch in echolocation phases is understandable. While flying with conspecifics, the lab's previous paper has shown what they call a 'clutter response' - but this is not necessarily the same as going into a 'buzz'-type call behaviour. As pointed out by another reviewer - the results of the simulations may hinge on the fact that bats are showing this echolocation phase-switching, and thus improving their echo-detection. This is not necessarily a major flaw - but something for readers to consider in light of the sparse experimental evidence at hand currently.

The use of echolocation phases—defined as the sequential search, approach, and buzz call patterns—has been documented not only during foraging but also in tasks such as landing, obstacle avoidance, clutter navigation, and drinking. Bat call structure has been shown to vary systematically with object proximity, not exclusively in response to prey. During obstacle avoidance, phase transitions were observed, with approach calls emitted in grouped sequences and with reduced durations (Gustafson & Schnitzler, 1979; Schnitzler et al., 1987). In landing contexts, bats have been reported to emit short-duration calls and decrease inter-pulse intervals—buzz-like patterns also observed during prey capture— suggesting shared acoustic strategies across behaviors (Hagino et al., 2007; Hiryu et al., 2008; Melcón et al., 2007, 2009). Comparable patterns have been reported during drinking maneuvers, where “drinking buzzes” have been proposed to guide a precise approach to the water surface, analogous to landing buzzes (Griffiths, 2013; Russo et al., 2016). In response to environmental complexity, bats were found to shorten calls and increase repetition rates when navigating cluttered spaces compared to open ones (Falk et al., 2014; Kalko & Schnitzler, 1993).

Moreover, field recordings from our study of Rhinopoma microphyllum (Goldshtein et al., 2025) revealed shortened call durations and inter-pulse intervals during dense group flight outside the cave during emergence—patterns consistent with terminal-approach phase that is typical when coming very close to an object (another bat in this case). The Author response image 1 shows an approach sequence recorded from a tagged bat approximately 20 meters from the cave entrance, with self-generated echolocation calls marked. The inter-pulse-interval of ca. 20 ms is used by these bats when a reflective object (another bat in this case) is nearby.

These results provide direct evidence that bats actively employ approach-phase echolocation during swarming likely to avoid collision with other bats. This supports the view that echolocation phase transitions are a general proximity-based sensing strategy, adapted across a variety of behavioral scenarios—not limited to hunting alone.

In our simulations, bats predominantly emitted calls in the approach phase, with only rare occurrences of buzz-phase calls.

See lines 355-363 in the revised manuscript.

The decision to model direction-of-arrival with such high angular resolution (1-2 degrees) is not entirely justifiable - and the authors may wish to do simulation runs with lower angular resolution. Past experimental paradigms haven't really separated out target-strength as a confounding factor for angular resolution (e.g. see the cited Simmons et al. 1983 paper). Moreover, to this reviewer's reading of the cited paper - it is not entirely clear how this experiment provides source-data to support the DoA-SNR parametrisation in this manuscript. The cited paper has two array-configurations, both of which are measured to have similar received levels upon ensonification. A relationship between angular resolution and signal-to-noise ratio is understandable perhaps - and one can formulate such a relationship, but here the reviewer asks that the origin/justification be made clear. On an independent line, also see the recent contrasting results of Geberl, Kugler, Wiegrebe 2019 (Curr. Biol.) - who suggest even poorer angular resolution in echolocation.

We thank the reviewer for raising this important point. The acuity of 1.5–3° in horizontal direction-of-arrival (DoA) estimation is based on the classical work of Simmons et al. with Eptesicus fuscus (Simmons et al., 1983). Similar precision was later supported by Erwin et al. (Erwin et al., 2001), who modeled azimuth estimation from measured interaural intensity differences (IIDs), reporting an average error of 0.2° with a standard deviation of ~2.2°, consistent with the behavioral data found by Simmons. The decline in acuity with increasing arrival angle has also been demonstrated in behavioral and physiological studies of binaural IID processing (Erwin et al., 2001; Fay, 1995; Razak, 2012; Wohlgemuth et al., 2016). The error model itself was first introduced in our earlier work (Mazar & Yovel, 2020).

Importantly, Geberl et al. (Geberl et al., 2019) examined the resolution of weak targets masked by nearby strong flankers and found poor spatial discrimination of ~45 degrees; however, they were studying a detection problem, rather than the horizontal acuity of azimuth estimation. Indeed, our model assumes there is no spatial discrimination at all.

Overall, while our DoA–SNR parametrization can certainly be critiqued and alternative parameterizations could be tested in future work, we believe it reflects a reasonable and empirically supported assumption.

Reviewer #2 (Public review):

This manuscript describes a detailed model for bats flying together through a fixed geometry. The model considers elements which are faithful to both bat biosonar production and reception and the acoustics governing how sound moves in air and interacts with obstacles. The model also incorporates behavioral patterns observed in bats, like one-dimensional feature following and temporal integration of cognitive maps. From a simulation study of the model and comparison of the results with the literature, the authors gain insight into how often bats may experience destructive interference of their acoustic signals and those of their peers, and how much such interference may actually negatively effect the groups' ability to navigate effectively. The authors use generalized linear models to test the significance of the effects they observe.

The work relies on a thoughtful and detailed model which faithfully incorporates salient features, such as acoustic elements like the filter for a biological receiver and temporal aggregation as a kind of memory in the system. At the same time, the authors abstract features that are complicating without being expected to give additional insights, as can be seen in the choice of a two-dimensional rather than three-dimensional system. I thought that the level of abstraction in the model was perfect, enough to demonstrate their results without needless details. The results are compelling and interesting, and the authors do a great job discussing them in the context of the biological literature.

With respect to the first version of the manuscript, the authors have remedied all my outstanding questions or concerns in the current version. The new supplementary figure 5 is especially helpful in understanding the geometry.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Data Availability: This reviewer lauds the authors for switching from a private commercial folder requiring login to one that does not. At the cost of being overtly pedantic - the Github repository is not a long-term archival resource. The ideal solution is to upload the code in an academic repository (Zenodo, OSF, etc.) to periodically create a 'static snapshot' of code for archival, while also hosting a 'live' version on Github.

We have uploaded to Zenodo repository, and updated the link in the paper:

How bats exit a crowded colony when relying on echolocation only - a modeling approach

In one of the rebuttals to Reviewer #3- the authors have cited a wrong paper (Beleyur & Goerlitz 2019) - while discussing broad bandwidth calls improving detection - and may wish to correct this if possible on record.

We have removed the incorrect citation from the revised version of the manuscript.

Specific comments on the 2nd manuscript:

Figure 5: Table 1 says 1, 2,5,10,20,40,100 bats were simulated (line 138-139) but the conclusion (line 398) says '1 to 100 bats' per 3msq. However, the X-axis only stops at 40 and says 'number of bats', while the legend says bats/3msq....what is actually being plotted? Moreover, in the entire paper there is a constant back-and-forth between density and # of bats - perhaps it is explained beforehand, but it is a bit unsettling - and more can be done to clarify these two conventions.

While most parameters were tested across the full range of 1 to 100 bats per 3 m², a subset of conditions—including misidentification, multi-call clustering, wall target strength, and conspecific target strength—were simulated only up to 40 bats due to significantly longer run-times. This is now clarified in both the main text and the Table 1 caption.

In our simulations, the primary parameter was the number of bats placed within a 3 m² starting area, which directly determined the initial density (bats per 3 m²). Throughout the manuscript, we use “number of bats” to refer to the simulation input, while “density” denotes the equivalent ecological measure. Figure 5 and related captions have been revised accordingly to note these conventions and to indicate when results are shown only up to 40 bats (see lines 120–122, 314-317 in the revised text).

Table 1: This was made considerably difficult to read given the visual clutter - and I hope I've understood these changes correctly.

What is in the square brackets of the effect-size (e.g. first row with values 'Exit prob. (%)' says -0.37/bat [63:100] ? What does this 63:100 refer to?

What is the 'process flag'

Values in square brackets indicate the minimum and maximum values of the metric across the tested range (e.g., [63:100] shows the range of exit probabilities observed across different bat densities).

The term “process flag” has been replaced with “with and without multi-call clustering” for clarity

Both the table layout and caption have been revised to reduce visual clutter and to make these conventions clearer to the reader.

Lines 562-3: "In our study, due to the dense cave environment, the bats are found to operate in the approach phase nearly all of the time, which is consistent with natural cave emergence behavior" - bats are 'found to' implies there is some experimental data or it is an emergent property. See above for the point questioing the implementation of multiple echolocation phases in the model, but also - here the bat-agents are allowed to show different phases and thus they do so -- it is a constraint of the implementation and not a result per se given the size of the cave and the number of bats involved...

We removed the sentence from the Methods section, since it could be misinterpreted as an experimental finding rather than a model outcome. Instead, we now discuss this in the Discussion, clarifying that the predominance of the approach phase arises from the cluttered cave environment in our simulations, which is consistent with natural emergence behavior (see lines 355-363). In this context, the use of echolocation phases is presented as a biologically plausible modeling choice rather than an empirical result.

Lines 659-660: The parametrisation between DoA and SNR is supposedly found in 'Equation 10' - which this reviewer could not find in the manuscript

The equation was accidentally omitted in the previous revision and has now been reinserted into the manuscript. It defines how direction-of-arrival (DoA) error depends on SNR and azimuth angle (see lines 603-605).

Changes from the earlier version to the revised version

Removed in the revised version: Abstract

Abstract

Bats face a complex navigation challenge when emerging from densely populated roosts, where vast numbers take off at once in dark, confined spaces. Each bat must avoid collisions with walls and conspecifics while locating the exit, all amidst overlapping acoustic signals. This crowded environment creates the risk of acoustic jamming, in which the calls of neighboring bats interfere with echo detection, potentially obscuring vital information. Despite these challenges, bats navigate these conditions with remarkable success. Although bats have access to multiple sensory cues, here we focused on whether echolocation alone could provide sufficient information for orientation under such high-interference conditions. To explore whether and how they manage this challenge, we developed a sensorimotor model that mimics the bats’ echolocation behavior under high-density conditions. Our model suggests that the problem of acoustic jamming may be less severe than previously assumed. Frequent calls with short inter-pulse intervals (IPI) increase the sensory input flow, allowing integration of echoic information across multiple calls. When combined with simple movement-guidance strategies—such as following walls and avoiding nearby obstacles—this accumulated information enables effective navigation in dense acoustic environments. Together, these findings demonstrate a plausible mechanism by which bats may overcome acoustic interference and underscore the role of signal redundancy in supporting robust echolocation-based navigation. Beyond advancing our understanding of bat behavior, they also offer valuable insights for swarm robotics and collective movement in complex environments.

Introduction

In many bat species individuals dwell together in caves (or similar roosts), forming large colonies with tens to several millions of individuals1,2. Each evening, at approximately the same time, the bats take off from their roost, navigating through its passages toward the exit. The high density of bats flying simultaneously in great proximity poses many challenges for orientation in such a crowded and noisy environment. Flying while avoiding collisions, often in a pitch-black cave, demands the continuous detection and localization of both obstacles and nearby bats3,4. Employing echolocation, bats emit strong ultrasonic signals and interpret the reflected echoes to perceive their surroundings5. The reception of neighbors’ loud calls, which share similar acoustic features with their own calls, can potentially hinder the bats’ ability to detect the faint echoes reflected off the walls and the surrounding bats5,6. We examined whether bats could rely solely on echolocation to exit the roost even during such a chaotic ‘rush hour’.

The question of how bats cope with acoustic interference — i.e., the masking of potential echoes by conspecific signals — has been extensively researched using playback experiments, field observations, on-body tags, and computational simulations7–17. However, much of this research has focused on foraging bats in small groups5,6,9,16,18–20. The challenges bats encounter during roost exits (e.g., cave exits) differ markedly from those encountered during group foraging. Bat density during roost exits is significantly higher, and bats need to detect and follow static walls or obstacles, which produce loud echoes, rather than small, sporadic prey items that generate faint echoes21. Their flight during exits is also more directional and involves avoiding collisions with conspecifics, in contrast to the erratic hunting maneuvers typically observed while foraging. Echolocation studies during dense collective movement are scarce4,6,22–25, likely due to the complexities in recording separate echolocation calls and tracking individual flights within the swarm.

While collective movement has been extensively studied in various species, such as insect swarming, fish schooling, and bird murmuration26–32, as well as in swarm robotics, where agents perform tasks such as coordinated navigation and maze-solving33–35, most studies have focused on movement algorithms that assume full detection of neighbors36–43. Some models have incorporated limited interaction rules where individuals respond to only one or a few neighbors due to sensory constraints44,45 However, fewer studies have explicitly examined how sensory interference, occlusion, and noise influence decision-making and affect collective movement46.

The present study addresses these gaps by introducing an agent-based sensorimotor model based on the well-documented echolocation capabilities of bats, simulating multiple bats pathfinding their way out of a cave-like structure. We modeled the echolocation behavior of two insectivorous bat species: Pipistrellus kuhlii (PK), which roosts in abandoned buildings and frequently navigates through conspecific-dense, cluttered corridors and the cave dwelling Rhinopoma microphyllum (RM) which emerges from its roosts with thousands of individuals simultaneously. These two species differ in their echolocation signals - PK echolocation signals are characterized by a wider bandwidth and a higher terminal frequency than RM calls. We quantified the performance of an individual bat flying among conspecifics, demonstrating that even a relatively simple sensorimotor algorithm can facilitate successful orientation in such complex environments. The modeling approach enabled us to explore how various biological and ecological factors influence successful navigation under such challenging conditions.

Results

Our model was designed with conservative assumptions regarding bats’ sensing, movement, and sensorimotor integration, aiming to underestimate their capabilities and thereby establish a lower bound on their actual performance. Real bats likely outperform the model’s predictions. In our 2D simulations7, each bat emits sound signals and receives echoes reflected from the roost walls and other bats, while also encountering masking signals caused by calls from conspecifics. These masking signals can interfere and completely eliminate echo detection (which we refer to as jamming) or cause echo localization errors. After estimating the distance and direction of each detected reflector, the bat adjusts its echolocation parameters and maneuvers to find the exit while simultaneously avoiding collisions. The bats dynamically adjust their echolocation parameters—including call rate, duration, and frequency range—based on the estimated distance to obstacles, following the well-documented transition between search, approach, and buzz phases observed in echolocating bats (see7 and Methods). Their reception was modeled using a biologically inspired filter-bank receiver comprising 80 gammatone channels7,47,48. Each bat adjusted its flight following a simple pathfinding algorithm based solely on the estimated locations of the detected reflectors (see Methods, Supplementary Figure 1, and Supplementary Movie 1 for additional details). The bats had to exit a roost designed as a corridor (14.5 m long x 2.5 m wide), with a right-angle turn located 5.5 m before the exit (Figure 1A). Additionally, an obstacle (1.25 m wide) was situated 2.25 m in front of the opening. The simulated bats initiated their flight from the far end of the corridor, within a randomly selected 1.5 × 2 m² area, taking off in the general direction of the exit (±30 degrees), without prior knowledge of the roost’s structure.

The sensorimotor model.

Earlier version: (A) Figure 2: The sensorimotor model. (A) Top view of the cave with three bats’ trajectories. The focal bat is shown in black. All bats’ flight trajectories are displayed while the bats’ moment-to-moment decisions are represented by the colored lines: blue – random flight, yellow – collision avoidance, light green – wall-following, turquoise – movement toward a wall gap (see panel D for details). Green squares depict reflectors detected by the focal bat along its route. (B) A zoomed-in view of the marked rectangular area in Panel A, where the focal bat (black) emitted one echolocation call (black) and received echoes from the cave walls (green) and from two other bats (blue). It also received conspecifics’ calls (red) and their reflection from the cave walls (orange), as well as the reflections from other bats (pink). Green squares indicate points that were detected by the focal bat from this call and red x’s indicate missed points due to acoustic masking (i.e., jammed reflectors). The locations of the detected reflectors (green squares) are marked according to their localization by the bat (with simulated errors). The lines near the bats depict their flight direction. (C) The acoustic scene received by the focal bat is as depicted in B, including the emitted call and all received signals (colors as in panel B). (C1) The time-domain plot displays the envelope of signals, encompassing the emitted call and the received signals: the desired echoes from the walls and conspecifics; the calls of other bats; the echoes returning from conspecific calls and reflected off the walls and off other bats. Notably, in this example, some of the desired wall-echoes are jammed by stronger self-echoes reflected from nearby conspecifics. (C2) The spectrogram of all the received signals presented in C1; for clarity, the emitted call is not depicted. (C3) The responses of the active channels of the cochlear filter bank (FB channel) after de-chirping. Each channel is represented by its central frequency on the y-axis. Each black dot represents the timing of a reaction that was above the detection threshold in each channel. Note that early reactions in low-frequency channels (marked by a yellow arrow) result from the stimulation of those channels caused by the higher frequencies of the downward FM chirp. However, most of these stimulations do not reach the detection threshold and are therefore not detected (see Methods). (C4) The detections of each channel are convolved with a Gaussian kernel, summed, and compared with the detection threshold (dotted red line). Colored asterisks mark peaks that were classified as successful detections—those identified in both the interference-free and full detection conditions (see Methods for details). Other peaks may originate from masking signals or overlapping echoes that did not meet the detection criteria (colors of the sources are as defined above). Panel D depicts the pathfinding algorithm used by the bat. The algorithm involves a correlated-random flight during the search phase (blue), collision avoidance (yellow), flying along the wall at a constant distance (green), and flying toward the center of a gap between obstacles as an indicator of a possible exit (cyan). After each echolocation call, the bat awaits an IPI (Inter Pulse Interval) period before processing the detections, adjusting flight and echolocation parameters, and emitting the next call. Based on the received signals, it then modifies its next call design and adjusts its direction and speed accordingly. For a detailed diagram of the complete sensorimotor process see Supplementary Figure 1.

Revised version: (A) Top view of the cave with three bats’ trajectories. The focal bat is shown in black. All bats’ flight trajectories are displayed while the bats’ moment-to-moment decisions are represented by the colored lines: blue - random flight, yellow – collision avoidance, light green – wall-following, turquoise – movement toward a wall gap (see panel D for details). Green squares depict reflectors detected by the focal bat along its route. (B) A zoomed-in view of the marked rectangular area in Panel A, where the focal bat (black) emitted one echolocation call (black) and received echoes from the cave walls (green) and from two other bats (blue). It also received conspecifics’ calls (red) and their reflection from the cave walls (orange), as well as the reflections from other bats (pink). Green squares indicate points that were detected by the focal bat from this call and red x’s indicate missed points due to acoustic masking (i.e., jammed reflectors). The locations of the detected reflectors (green squares) are marked according to their localization by the bat (with simulated errors). The lines near the bats depict their flight direction. (C) The acoustic scene received by the focal bat is as depicted in B, including the emitted call and all received signals (colors as in panel B). (C1) The time-domain plot displays the envelope of signals, encompassing the emitted call and the received signals: the desired echoes from the walls and conspecifics; the calls of other bats; the echoes returning from conspecific calls and reflected off the walls and off other bats. Notably, in this example, some of the desired wall-echoes are jammed by stronger self-echoes reflected from nearby conspecifics. (C2) The spectrogram of all the received signals presented in C1; for clarity, the emitted call is not depicted. (C3) The responses of the active channels of the cochlear filter bank (FB channel) after de-chirping. Each channel is represented by its central frequency on the y-axis. Each black dot represents the timing of a reaction that was above the detection threshold in each channel. Note that early reactions in low-frequency channels (marked by a yellow arrow) result from the stimulation of those channels caused by the higher frequencies of the downward FM chirp. However, most of these stimulations do not reach the detection threshold and are therefore not detected (see Methods). (C4) The detections of each channel are convolved with a Gaussian kernel, summed, and compared with the detection threshold (dotted red line). Colored asterisks mark peaks that were classified as successful detections—those identified in both the interference-free and full detection conditions (see Methods for details). Other peaks may originate from masking signals or overlapping echoes that did not meet the detection criteria (colors of the sources are as defined above). Panel D depicts the pathfinding algorithm used by the bat. The algorithm involves a correlated-random flight during the search phase (blue), collision avoidance (yellow), flying along the wall at a constant distance (green), and flying toward the center of a gap between obstacles as an indicator of a possible exit (cyan). After each echolocation call, the bat awaits an IPI (Inter Pulse Interval) period before processing the detections, adjusting flight and echolocation parameters, and emitting the next call. Based on the received signals, it then modifies its next call design and adjusts its direction and speed accordingly. For a detailed diagram of the complete sensorimotor process see Supplementary Figure 1.

The sensory model accounted for six types of acoustic signals: (1) the bat’s own calls, (2) echoes from conspecifics, (3) echoes from walls in response to the bat’s own calls (i.e., desired wall echoes), (4) echoes from conspecific calls reflected off other bats, (5) echoes from conspecific calls reflected off walls, and (6) the conspecific calls themselves. In the baseline model, bats were assumed to reliably distinguish between all these signal types. In contrast, the confusion model described below specifically tested the impact of failing to distinguish between desired wall echoes (3) and wall echoes generated by conspecific calls (5), while preserving the bat’s ability to identify all other signal types. In brief, the bat responded to echoes as follows (see Methods and Supplementary Figure 1 for details): If an obstacle or a conspecific was detected in front of the bat and was too close, the bat would maneuver to avoid a collision. Otherwise, for exit-seeking, the bat would follow the contour of the walls by steering toward the farthest detected obstacle ahead. If a gap greater than 0.5 m was identified between adjacent reflectors, the bat directed its trajectory toward the center of the gap.

Earlier version: The ability of the bats to exit the roost within 15 sec was evaluated for different group sizes, from a single bat and up to 100 individuals. For simplicity, we will refer to the initial density at the cave’s far end as the number of bats per 3m2 (i.e., for groups of 100 bats, the density is 100 bats/3m2, or 33.3 bats/m2). The bat densities we tested were chosen to reflect the typical range of bat densities observed in natural caves during emergence events25,49,50. Key model parameters, such as the sensory integration window, object target strength, echolocation parameters, and flight velocity (see Table 1), were manipulated and their impact on the exit performance was analyzed. To explicitly quantify the effect of sensory masking vs. the effect of collision avoidance (i.e., spatial interference) only, we turned the acoustic interference on and off to measure its impact. The number of repetitions for each scenario was as follows: 1 bat: 240; 2 bats: 120; 5 bats: 48; 10 bats: 24; 20 bats: 12; 40 bats: 12; 100 bats: 6 (see Table 1).

Revised version: The ability of the bats to exit the roost within 15 sec was evaluated for different group sizes, from a single bat and up to 100 individuals. For simplicity, we will refer to the initial density at the cave’s far end as the number of bats per 3m2 (i.e., for groups of 100 bats, the density is 100 bats/3m2, or 33.3 bats/m2). The bat densities we tested were chosen to reflect the typical range of bat densities observed in natural caves during emergence events25,49,50. Key model parameters, such as the sensory integration window, object target strength, echolocation parameters, and flight velocity (see Table 1), were manipulated and their impact on the exit performance was analyzed. To explicitly quantify the effect of sensory masking vs. the effect of collision avoidance (i.e., spatial interference) only, we turned the acoustic interference on and off to measure its impact. Each scenario was repeated as follows: 1 bat: 240; 2 bats: 120; 5 bats: 48; 10 bats: 24; 20 bats: 12; 40 bats: 12; 100 bats: 6 (see Table 1), while misidentification rate, multi-call clustering, and wall/conspecific target strength were tested only up to 40 bats (see Table 1).

Bats find their way out of the cave even at high conspecific densities

Earlier version: We first examined how bat density affects bats’ ability to exit the cave, both alone and in a group. The probability of exiting the cave within 15 seconds—defined as the proportion of bats that successfully exited within this time frame—was significantly reduced at higher densities (Figure 3A, see Supplementary Movie 1 for a view of the bats’ movement, p<10-10, t =-23, DF=4077, GLM, see details in Table 1). In trials in which a single bat was flying alone, it successfully exited the cave in 100% of the cases. Even without sensory interference, the probability of exiting decreased significantly from 100% to 86%±1.4% and 91%±1.7% at densities of 100 PKs/3m2 and 100 RMs/3m2, respectively (mean ± s.e.). When acoustic interference was added, the exit probability further decreased to 63%±1.4% and 67%±1.4% for 100 PKs and RMs, respectively (see Figure 3A).

Revised version: We first examined how bat density affects bats’ ability to exit the cave, both alone and in a group. The probability of exiting the cave within 15 seconds—defined as the proportion of bats that successfully exited within this time frame—was significantly reduced at higher densities (Figure 2A, see Supplementary Movie 1 for a view of the bats’ movement, p<10-10, t =-23, DF=4077, GLM, see details in Table 1). In trials in which a single bat was flying alone, it successfully exited the cave in 100% of the cases. Even without sensory interference, the probability of exiting decreased significantly from 100% to 86%±1.4% and 91%±1.7% at densities of 100 PKs/3m2 and 100 RMs/3m2, respectively (mean ± s.e.). When acoustic interference was added, the exit probability further decreased to 63%±1.4% and 67%±1.4% for 100 PKs and RMs, respectively (see Figure 2A).

Exit performance of P. Kuhlii (PK) and R. Microphyllus (RM).

(A) Sensory interference significantly impaired the probability of exiting the cave (compare dashed lines with continuous lines). The probability of a successful exit also declined as the number of bats increased, with no significant difference observed between the species when masking interference was applied. The insert shows the spectrograms of the echolocation calls of PK (top) and RM (bottom). (B) The time-to-exit, which was calculated for successful trials only, and (C) the collision rate with the walls both increased as a function of the number of bats. (D) The probability of jamming significantly increased to about 55% and 63% with 100 bats for PK and RM, respectively. (E) The detection probability of a wall reflector at one meter or less in front of a bat decreased as a function of the number of bats. In panels (A-E), circles represent means and bars represent standard errors (see details in Table 1). Asterisks indicate significant differences between the lines in each panel.

Earlier version: The difference in exit probability between the two species was not significant (p=0.08, t =1.74, DF=4077, GLM as above, Figure 3A). Similarly, the difference in echolocation parameters between the two species did not affect the collision rate with the walls (with a maximum of 0.29 and 0.3 collisions per bat per second for PK and RM, respectively, with 100 bats (p=0.63, t =-0.48, DF=4077, GLM, Figure 3C, see details in Table 1). To quantify sensory interference, we defined a jammed echo as an echo entirely missed due to masking. The jamming probability, which was calculated as the number of jammed echoes divided by the total number of self-echoes, was significantly higher for RM compared to PK. The maximum difference between the two models was 14.3% at a density of 10 bats, with a smaller difference of 9.8% observed at 100 bats (p<10- 10, t =6.56, DF=4077, GLM, Figure 3D, see details in Table 1). Accordingly, PK demonstrated a minor but significant advantage in detecting the cave walls (p=0.024, t =-2.25, DF=4077, GLM, Figure 3E, see details in Table 1). With 100 bats flying together, the probability of detecting a wall echo at a distance of 1 m in a single call was around 50% and 46% for PK and RM, respectively. Despite this minor disadvantage in detection, RM bats exhibited a better time-to-exit average than PK bats, being 0.5 seconds faster to exit (p=0.0005, t =-4.06, DF=3533, for n=40 bats, Figure 3B). Additionally, RM bats experienced a significantly higher probability of their self-generated echoes, reflected off conspecifics, being jammed (p = 0.00016, t = 3.8, DF = 3593, GLM; see details in Table 1).

Revised version: The difference in exit probability between the two species was not significant (p=0.08, t =1.74, DF=4077, GLM as above, Figure 2A). Similarly, the difference in echolocation parameters between the two species did not affect the collision rate with the walls (with a maximum of 0.29 and 0.3 collisions per bat per second for PK and RM, respectively, with 100 bats (p=0.63, t =-0.48, DF=4077, GLM, Figure 2C, see details in Table 1). To quantify sensory interference, we defined a jammed echo as an echo entirely missed due to masking. The jamming probability, which was calculated as the number of jammed echoes divided by the total number of self-echoes, was significantly higher for RM compared to PK. The maximum difference between the two models was 14.3% at a density of 10 bats, with a smaller difference of 9.8% observed at 100 bats (p<10-10, t =6.56, DF=4077, GLM, Figure 2D, see details in Table 1). Accordingly, PK demonstrated a minor but significant advantage in detecting the cave walls (p=0.024, t =-2.25, DF=4077, GLM, Figure 2E, see details in Table 1). With 100 bats flying together, the probability of detecting a wall echo at a distance of 1 m in a single call was around 50% and 46% for PK and RM, respectively. Despite this minor disadvantage in detection, RM bats exhibited a better time-to-exit average than PK bats, being 0.5 seconds faster to exit (p=0.0005, t =-4.06, DF=3533, for n=40 bats, Figure 2B). Additionally, RM bats experienced a significantly higher probability of their self-generated echoes, reflected off conspecifics, being jammed (p = 0.00016, t = 3.8, DF = 3593, GLM; see details in Table 1).

Multi-call integration improves exit performance

Earlier version: We next examined whether bats improve their performance when integrating information from several consecutive calls. The integration window determines the number of previous calls the bat uses at each step to guide its next movement decision (see Methods and Supplementary Figure 2A). In the basic multi-call integration model, detections from the previous calls — by default the last five — were stored in an allocentric (x-y) reference frame, with each detection treated independently as a potential obstacle without clustering or filtering. At each decision, the bat takes all of these detections into account when guiding its movement and echolocation. The probability of exiting the roost significantly increased when increasing the size of the integration window for all bat densities (p<10-10, t =28.5, DF=10197, GLM, Figure 4A, see details in Table 1). For example, at a density of 40 bats/3m2, the exit probability improved from 20%, to 75%, and to 87% as the window size increased from one, to three, and to 10 previous calls, respectively. In addition, increasing the window size resulted in a significant improvement in the time-to-exit and the avoidance of wall collisions (p<10-10, t =-12.8, DF=7661; p<10-10, t =-46.5, DF=10197, respectively, GLM, see details in Table 1). With 100 bats, the collision rate decreased by a factor of 2 from 0.53 to 0.25 collisions per second as the window increased from 1 to 10 calls. The size of the integration window had no significant effect on the jamming probability (p=0.37, t =0.9, DF=10197, GLM, see details in Table 1).

Revised version: We next examined whether bats improve their performance when integrating information from several consecutive calls. The integration window determines the number of previous calls the bat uses at each step to guide its next movement decision (see Methods and Supplementary Figure 2A). In the basic multi-call integration model, detections from the previous calls — by default the last five — were stored in an allocentric (x-y) reference frame, with each detection treated independently as a potential obstacle without clustering or filtering. At each decision, the bat takes all of these detections into account when guiding its movement and echolocation. The probability of exiting the roost significantly increased when increasing the size of the integration window for all bat densities (p<10-10, t =28.5, DF=10197, GLM, Figure 3A, see details in Table 1). For example, at a density of 40 bats/3m2, the exit probability improved from 20%, to 75%, and to 87% as the window size increased from one, to three, and to 10 previous calls, respectively. In addition, increasing the window size resulted in a significant improvement in the time-to-exit and the avoidance of wall collisions (p<10-10, t =-12.8, DF=7661; p<10-10, t =-46.5, DF=10197, respectively, GLM, see details in Table 1). With 100 bats, the collision rate decreased by a factor of 2 from 0.53 to 0.25 collisions per second as the window increased from 1 to 10 calls. The size of the integration window had no significant effect on the jamming probability (p=0.37, t =0.9, DF=10197, GLM, see details in Table 1).

Exit performance as a function of key sensorimotor parameters.

(A) The effect of the integration window on the probability of exiting the roost, the time-to exit, the rate of collisions with the walls, and the probability of jamming (from left to right, respectively). Each colored line shows the trend as a function of the window-size for different bat densities, with each color representing a specific density. Note that a window size of 0 indicates that only the most recent call is used in the bat’s decision-making, without integrating detections from previous calls. (B) The effect of the nominal flight speed of the bats, with panels and line-colors as in panel A. An optimal speed of approximately 6 to 8 m/sec can be observed for all densities above one bat. (C) The effect of call intensity on exit performance, panels as in (A). In all panels, circles represent means and bars represent standard errors. Error bars depicting standard errors are presented but are very small due to the large number of simulation repetitions. See Table 1 for the number of simulated bats.

Exit probability was maximal at an intermediate flight-speed

Earlier version: We observed a significant and non-linear effect of the flight speed of the bats on the performance, as shown in Figure 3B (p<10-10, t =-29.9, DF=10196, GLM, see details in Table 1). The exit probability increased with flight speed until it reached a maximum at 6-8 m/s and then declined rapidly. This was the case for all bat densities, with the maximal exit probability ranging between 65% to 99%. At the optimal velocity, the time-to-exit was also minimal. However, the collision rate increased monotonically with speed, with a steep incline above the optimal speed.

Revised version: We observed a significant and non-linear effect of the flight speed of the bats on the performance, as shown in Figure 1 Figure 3B (p<10-10, t =-29.9, DF=10196, GLM, see details in Table 1). The exit probability increased with flight speed until it reached a maximum at 6-8 m/s and then declined rapidly. This was the case for all bat densities, with the maximal exit probability ranging between 65% to 99%. At the optimal velocity, the time-to-exit was also minimal. However, the collision rate increased monotonically with speed, with a steep incline above the optimal speed.

Call intensity had only a minor effect on exit performance and only at high bat densities

Earlier version: For low bat densities (<40 bats), call intensity did not have a significant impact on either exit probability or collision rate (Figure 4C, p=0.89, t =0.13, DF=5757; p=82, t –0.21, DF=5757, respectively, GLM, see details in Table 1). Call intensity affected exit performance only when the intensity dropped to 100 dB-SPL (@ 0.1m) and only at a high bat density of 100 bats/3m2 (Figure 4 C). In this scenario, the exit probability declined from approximately 60% to 49.5% (p=0.003, F = 8.45, DF = 2396, One-way ANOVA with ‘hsd’ post hoc test), and the collision rate increased from 0.3 to 0.35 collisions per second (p<3·10-6, F = 22.18, DF = 2396). Notably, this low intensity is below the typical search-call intensity of most echolocating bats. At the same bat density (100 bats/3m2), further increasing the call intensity to above 100dB-SPL had no significant effect on either exit probability (p=0.6) or collision rate (p=0.07). Calling louder also slightly, but significantly, decreased the jamming probability at all bat densities, with a decrease of 3.5%±8% to 5.5%±5% (mean ± s.e.) (p=0.02, t =-2.26, DF=8157, GLM, see Table 1).

Revised version: For low bat densities (<40 bats), call intensity did not have a significant impact on either exit probability or collision rate (Figure 3C, p=0.89, t =0.13, DF=5757; p=82, t -0.21, DF=5757, respectively, GLM, see details in Table 1). Call intensity affected exit performance only when the intensity dropped to 100 dB-SPL (@ 0.1m) and only at a high bat density of 100 bats/3m2 (Figure 3Figure 3 C). In this scenario, the exit probability declined from approximately 60% to 49.5% (p=0.003, F = 8.45, DF = 2396, One-way ANOVA with ’hsd’ post hoc test), and the collision rate increased from 0.3 to 0.35 collisions per second (p<3·10-6, F = 22.18, DF = 2396). Notably, this low intensity is below the typical search-call intensity of most echolocating bats. At the same bat density (100 bats/3m2), further increasing the call intensity to above 100dB-SPL had no significant effect on either exit probability (p=0.6) or collision rate (p=0.07). Calling louder also slightly, but significantly, decreased the jamming probability at all bat densities, with a decrease of 3.5%±8% to 5.5%±5% (mean ± s.e.) (p=0.02, t =-2.26, DF=8157, GLM, see Table 1).

While confusion between the desired echoes and those from conspecific calls may significantly impair exit performance, multi-call clustering helps to mitigate this

Earlier version: We next addressed the challenge of echo classification, assuming that a bat can differentiate an echo resulting from its own calls from echoes resulting from the calls of other bats. To examine this assumption, we tested another model, referred to as the confusion model, in which bats responded similarly both to wall echoes returning from their own emissions and to those from conspecific emissions, treating all as their own echoes. This confusion significantly decreased exit performance for all bat densities (above one bat). The probability of a successful exit for a density of 40 bats/3m2 dropped from 83.3±2.4% to 14.6±2.3% (p<<10-10, t =-20.7, DF=2877, GLM, see details in Table 1), the exit time increased from 7.6±0.18 to 9.3±0.2 seconds (p<<10-10, t =15.5, DF=2157, GLM), and the collision rate increased significantly from 0.2±0.007 to 0.8±0.013 collisions per second (p<<10-10, t =-30, DF=28777, GLM, see Figure 5, red and yellow lines).

Revised version: We next addressed the challenge of echo classification, assuming that a bat can differentiate an echo resulting from its own calls from echoes resulting from the calls of other bats. To examine this assumption, we tested another model, referred to as the confusion model, in which bats responded similarly both to wall echoes returning from their own emissions and to those from conspecific emissions, treating all as their own echoes. This confusion significantly decreased exit performance for all bat densities (above one bat). The probability of a successful exit for a density of 40 bats/3m2 dropped from 83.3±2.4% to 14.6±2.3% (p<<10-10, t =-20.7, DF=2877, GLM, see details in Table 1), the exit time increased from 7.6±0.18 to 9.3±0.2 seconds (p<<10-10, t =15.5, DF=2157, GLM), and the collision rate increased significantly from 0.2±0.007 to 0.8±0.013 collisions per second (p<<10-10, t =-30, DF=28777, GLM, see Figure 4, red and yellow lines).

The impact of confusion on performance.

The figure illustrates the impact of classification confusion on roost-exit performance under various conditions. Blue lines depict trials with masking, while assuming that bats can distinguish between echoes from their own calls and those of conspecifics (referred to as "No Confusion"). Red lines depict performance where confusion between echoes is assumed. Yellow lines depict performance under the confusion condition, with the added capability of multi-call clustering in a short-term working memory (referred to as “confusion with mitigation”, see text for further details). In all panels, circles represent means and bars indicate standard errors. (A) The probability of exiting the roost significantly decreased with masking and confusion. In conditions with confusion and no aggregation process, only 15% of bats successfully exited the roost, at a density of 40 bats/3m2. Multi-call clustering partially mitigated the confusion effect but did not eliminate it. (B) Bats with the ability to distinguish between echoes demonstrated significantly shorter exit times than those experiencing confusion. Note that time-to-exit refers only to successful attempts. (C) The collision rate with walls was highest for bats experiencing both masking and confusion but decreased significantly when without confusion. Multi-call clustering restored performance to the "No Confusion" condition, reducing collision rates accordingly, at densities between 1 to 40 bats/3m2.

Earlier version: To further examine whether this substantial decrease in performance could be mitigated even without improving echo identification, we tested an enhanced integration model that, in addition to extending the number of calls integrated, clustered spatially close detections, removed outliers, and estimated wall directions based on grouped reflectors (see Methods and Supplementary Figure 2B). This ’multi-call clustering’ significantly improved performance, but exit probability and time-to-exit still remained significantly lower than without echo-confusion: exit probability = 58±3% in comparison to 83.3±2.4% without echo confusion (p<<10-10, t =18.3, DF=28777, GLM), time-to-exit =9.3±0.2 seconds (p<<10-10, t =-13.7, DF=1996, GLM), see Figure 4, yellow line. The results above are reported for a density of 40 bats/3m2. Interestingly, the multi-call clustering restored the collision rate to the levels observed under the “No Confusion” condition (p=0.68, t =-0.42, DF= 2877, GLM, see Figure 5C, dark-purple and red lines).

Revised version: To further examine whether this substantial decrease in performance could be mitigated even without improving echo identification, we tested an enhanced integration model that, in addition to extending the number of calls integrated, clustered spatially close detections, removed outliers, and estimated wall directions based on grouped reflectors (see Methods and Supplementary Figure 2B). This ’multi-call clustering’ significantly improved performance, but exit probability and time-to-exit still remained significantly lower than without echo-confusion: exit probability = 58±3% in comparison to 83.3±2.4% without echo confusion (p<<10-10, t =18.3, DF=28777, GLM), time-to-exit =9.3±0.2 seconds (p<<10-10, t =-13.7, DF=1996, GLM), see Figure 4, yellow line. The results above are reported for a density of 40 bats/3m2. Interestingly, the multi-call clustering restored the collision rate to the levels observed under the "No Confusion" condition (p=0.68, t =-0.42, DF= 2877, GLM, see Figure 4C, dark-purple and red lines).

Effect of Wall and Conspecific Target Strengths on Exit Performance

Increasing the wall target strength significantly enhanced navigation performance (Supplementary Figure 3, Table 1), improving exit probability by up to 64% and reducing time-to-exit by up to 2.8 seconds (p << 10⁻¹⁰). Stronger wall echoes improved environmental awareness but also slightly increased masking of desired conspecific signals.

In contrast, changes in conspecific target strength had a much smaller effect (Supplementary Figure 4, Table 1), with only minor improvements in detection and collision rates, and no significant impact on exit probability. This likely reflects the fact that both desired and masking signals scale similarly with conspecific reflectivity. Overall, the model showed low sensitivity to variations in conspecific target strength.

Discussion

We present a model-based approach that suggests how echolocating bats might find their way out of a crowded roost while contending with severe sensory interference caused by numerous nearby conspecifics. Our results demonstrate that a single bat, lacking prior knowledge of the roost’s structure, successfully found the exit in all simulated trials using echolocation alone. As bat density increases, the bats face increased collision risks and more substantial acoustic interference, both of which reduce the probability of efficiently finding the exit. Nevertheless, even at densities of 100 bats/3m2, most bats (63%) successfully exited the roost within a short timeframe. These results are based on a 2D simulation with up to 33 bats/m², under the assumption that bats can distinguish their own echoes from those of conspecifics. We demonstrate how a simple sensorimotor approach can solve this supposedly challenging task. This approach encompasses the following principles: (1) emission of echolocation calls; (2) reception of reflected echoes and masking signals; (3) detection of reflectors (including walls and conspecifics) using a gammatone filter bank biological receiver; (4) localization of the detected objects; (5) employment of multi-call integration of acoustic detections; (6) adjustment of flight and echolocation behavior based on the distance and angle to the reflectors; and (7) application of simple pathfinding rules to follow walls and gaps while avoiding collisions. Notably, despite the jamming of a substantial percentage of the echoes — particularly, with 100 bats, 50% of the echoes from nearby obstacles at ∼1 m distance — the bats managed to maneuver correctly even with this simple approach and partial data.

A key component of this success was the multi-call integration: increasing the number of stored calls from one to ten markedly improved performance, raising the exit probability from 20% to 87% and halving the collision rate. Real bats likely use a much more sophisticated approach that also includes memorizing the roost’s structure51, using landmarks inside the roost52, reliance on the movement of nearby conspecifics43,49, and exploitation of other sensory modalities. We thus expect their actual performance to surpass that of our modeled bats.

Our model suggests that acoustic jamming might be less problematic than has been generally assumed5,11,53, and that movement under severe acoustic masking could be mitigated by increasing the call-rate, creating a redundancy across several calls-similar to how real bats behave in a complex environment6. In our model, the Inter-Pulse Interval (IPI) naturally varied according to established echolocation behavior, decreasing from 100 msec in the search phase to 35 msec (∼28 calls per second) in the approach phase, and further to 5 msec (200 calls per second) during the final buzz (Table 2). The results indicate that this redundancy, combined with simple sensorimotor heuristics, enhances successful navigation. This is consistent with several recent studies that have pointed in this direction7,24,25.

Added in the revised version: While echolocation phases—search, approach, and buzz—are traditionally associated with prey capture, similar patterns have been documented in non-foraging tasks such as landing, obstacle avoidance, clutter navigation, and drinking54–64. In these contexts, bats modulate call duration and inter-pulse intervals according to object proximity, generating phase-like transitions even without prey. This supports the interpretation of phase structure as a general proximity-sensing strategy rather than a foraging-specific behavior. In our simulations, bats operated predominantly in the approach phase due to the cluttered cave environment—consistent with natural emergence behavior, where navigation dominates over open-space search. Accordingly, our use of echolocation phases in the model is biologically plausible across a range of sensory-guided tasks.

Earlier version: The bat densities we simulated, ranging from 1 to 100 bats per 3m2, reflect a wide range reported in field studies. Although bat colonies can be much larger than 100 bats, the maximal simulated density in our model (100 bats per 3 m²) resulted in bats flying in very close proximity, with an average nearest-neighbor distance of 0.27 meters. This density is higher than some of the most-dense reported bat aggregations, including studies on Miniopterus fuliginosus49, Myotis grisescens54, and Tadarida brasiliensis4,50,55, where bats emerge from the roost at rates of 15 to 500 bats per second, but fly with an average distance of 0.5 meters between individual bats.

Revised version: The bat densities we simulated, ranging from 1 to 100 bats per 3m2, reflect a wide range reported in field studies. Although bat colonies can be much larger than 100 bats, the maximal simulated density in our model (100 bats per 3 m²) resulted in bats flying in very close proximity, with an average nearest-neighbor distance of 0.27 meters. This density is higher than some of the most-dense reported bat aggregations, including studies on Miniopterus fuliginosus49, Myotis grisescens65, and Tadarida brasiliensis4,50,66, where bats emerge from the roost at rates of 15 to 500 bats per second, but fly with an average distance of 0.5 meters between individual bats.

Earlier version: We compared the performance of two FM echolocating insectivorous bat species: Pipistrellus kuhlii (PK) and Rhinopoma microphyllum (RM). PK bats emit wideband echolocation signals that are less prone to jamming than RM bats’ narrowband signal15,56, as wideband signals distribute energy across a broader frequency range and are thus more robust against interference9,24,57. Our findings show that PK signals slightly reduce jamming probability (by 9%) and improve wall detection. However, no significant differences in exit probabilities were noted between the two species.

Revised version: We compared the performance of two FM echolocating insectivorous bat species: Pipistrellus kuhlii (PK) and Rhinopoma microphyllum (RM). PK bats emit wideband echolocation signals that are less prone to jamming than RM bats’ narrowband signal15,67, as wideband signals distribute energy across a broader frequency range and are thus more robust against interference9,68. Our findings show that PK signals slightly reduce jamming probability (by 9%) and improve wall detection. However, no significant differences in exit probabilities were noted between the two species.

Earlier version: Using a simulation allowed us to separate the effects of acoustic interference (masking) and spatial interference (collision avoidance) and revealed new insights into the sensorimotor strategy that that could plausibly be used by real bats. The spatial interference reduced the probability of exiting the roost from 100% to 87%, while the acoustic masking further decreased it to 63%. Increasing call intensity had little effect on exit performance, although slightly improving it at high bat densities. When all bats increased their calling intensity, both desired echoes and masking signals intensified equally, resulting in only a marginal effect. This was tested by varying call intensity levels (100-130 dB SPL) in our simulations (Table 1), demonstrating that beyond a certain level (∼110 dB SPL), there is no further benefit in improving obstacle detection. These results align with previous studies that have drawn similar conclusions7,24.

Revised version: Using a simulation allowed us to separate the effects of acoustic interference (masking) and spatial interference (collision avoidance) and revealed new insights into the sensorimotor strategy that could plausibly be used by real bats. The spatial interference reduced the probability of exiting the roost from 100% to 87%, while the acoustic masking further decreased it to 63%. Increasing call intensity had little effect on exit performance, although slightly improving it at high bat densities. When all bats increased their calling intensity, both desired echoes and masking signals intensified equally, resulting in only a marginal effect. This was tested by varying call intensity levels (100-130 dB SPL) in our simulations (Table 1), demonstrating that beyond a certain level (∼110 dB SPL), there is no further benefit in improving obstacle detection. These results align with previous studies that have drawn similar conclusions7,24.

Earlier version: Bats constantly adjust their flight speed to their surroundings58–61 and specifically when conspecifics are nearby62. Our study suggests that the optimal velocity for flying through a crowded roost ranges from 6 m/sec to 8 m/sec for densities of 2-100 bats/3m2. Exceeding this velocity-range led to a significant drop in exit probability due to a significant increase in wall collisions. We found that this speed did not depend on bat density in accordance with the observations of Theriault et al.50. Notably, the reported velocities of RM when exiting a cave25 and PK emergence velocity near the cave63 are close to the speed that appears optimal, based on our simulations.

Revised version: Bats constantly adjust their flight speed to their surroundings69–72 and specifically when conspecifics are nearby73. Our study suggests that the optimal velocity for flying through a crowded roost ranges from 6 m/sec to 8 m/sec for densities of 2-100 bats/3m2. Exceeding this velocity-range led to a significant drop in exit probability due to a significant increase in wall collisions. We found that this speed did not depend on bat density in accordance with the observations of Theriault et al.50. Notably, the reported velocities of RM when exiting a cave25 and PK emergence velocity near the cave74 are close to the speed that appears optimal, based on our simulations.

We also tested the effects of wall and conspecific target strengths on navigation. Stronger wall echoes substantially improved exit probability and reduced obstacle collisions, despite slightly increasing masking of conspecific echoes (Supplementary Figure 3). In contrast, changes in conspecific reflectivity had minimal impact, likely because both desired and masking signals scaled similarly (Supplementary Figure 4). This result may also stem from our model’s assumption that bats slow down, but continue flying at the same direction following a collision with a conspecific.

Earlier version: Our basic model assumed that bats can distinguish between wall echoes and conspecific echoes, as demonstrated in previous studies 64–66. We suggest that this is a feasible assumption because echoes from cave walls are longer and exhibit distinct spectro-temporal patterns, whereas echoes from smaller objects, such as conspecifics, are shorter47,67,68. However, wall echoes reflected from conspecific calls might resemble those from the bat’s own calls in their amplitude and time-frequency characteristics 20,62,69. This led us to question how the misidentification of such echoes as obstacles might affect navigation. When unable to distinguish between these echoes, the simulated bats responded to all as if they were their own and thus mis-localized conspecific wall echoes. The confusion led to a substantial drop in exit performance, with only 15% of the bats successfully exiting compared to 82% under no-confusion conditions, at a density of 40 bats/3m2. At the same time, the collision rate increased markedly from 0.2 to 0.85 collisions per second. These results demonstrate the vital importance of echo discrimination for successful navigation, highlighting both the necessity of distinguishing between self and conspecific echoes and the classic challenge of detecting desired signals in noisy environments. There is a substantial evidence in the literature supporting the assumption that bats can recognize their own echoes and reliably distinguish them from those of conspecifics57,64–66,70.

Revised version: Our basic model assumed that bats can distinguish between wall echoes and conspecific echoes, as demonstrated in previous studies 75–77. We suggest that this is a feasible assumption because echoes from cave walls are longer and exhibit distinct spectro-temporal patterns, whereas echoes from smaller objects, such as conspecifics, are shorter47,78,79. However, wall echoes reflected from conspecific calls might resemble those from the bat’s own calls in their amplitude and time-frequency characteristics 20,73,80. This led us to question how the misidentification of such echoes as obstacles might affect navigation. When unable to distinguish between these echoes, the simulated bats responded to all as if they were their own and thus mis-localized conspecific wall echoes. The confusion led to a substantial drop in exit performance, with only 15% of the bats successfully exiting compared to 82% under no-confusion conditions, at a density of 40 bats/3m2. At the same time, the collision rate increased markedly from 0.2 to 0.85 collisions per second. These results demonstrate the vital importance of echo discrimination for successful navigation, highlighting both the necessity of distinguishing between self and conspecific echoes and the classic challenge of detecting desired signals in noisy environments. There is a substantial evidence in the literature supporting the assumption that bats can recognize their own echoes and reliably distinguish them from those of conspecifics68,75–77,81.

Earlier version: Previous studies have also demonstrated that bats can aggregate acoustic information received sequentially over several echolocation calls, effectively constructing an auditory scene in complex environments5,71–75. Bats are also known to emit call sequences in groups, particularly when spatiotemporal localization demands are high. Studies have recorded sequences of 2-15 grouped calls, supporting the idea that grouping facilitates echo aggregation72,76. Accordingly, we tested how multi-call clustering process—which included grouping nearby reflectors, removing outliers, and estimating wall orientation based on these clusters—could assist bats in pathfinding, even under the assumption of full confusion. At bat densities of 1 to 40 bats/3m2 with masking, the multi-call clustering completely restored the collision rate with walls from 0.85 back to 0.2 collisions per second, and significantly improved the exit probability, raising it to 58%, although it did not entirely eliminate the impact of confusion. Our assumption of total confusion between echoes from a bat’s own calls and those from conspecifics, as well as our relatively simple clustering model, likely underestimates the true capabilities of real bats when flying in complex environments.

Revised version: Previous studies have also demonstrated that bats can aggregate acoustic information received sequentially over several echolocation calls, effectively constructing an auditory scene in complex environments5,82–86. Bats are also known to emit call sequences in groups, particularly when spatiotemporal localization demands are high. Studies have recorded sequences of 2-15 grouped calls, supporting the idea that grouping facilitates echo aggregation83,87. Accordingly, we tested how multi-call clustering process—which included grouping nearby reflectors, removing outliers, and estimating wall orientation based on these clusters—could assist bats in pathfinding, even under the assumption of full confusion. At bat densities of 1 to 40 bats/3m2 with masking, the multi-call clustering completely restored the collision rate with walls from 0.85 back to 0.2 collisions per second, and significantly improved the exit probability, raising it to 58%, although it did not entirely eliminate the impact of confusion. Our assumption of total confusion between echoes from a bat’s own calls and those from conspecifics, as well as our relatively simple clustering model, likely underestimates the true capabilities of real bats when flying in complex environments.

Earlier version: Navigation in bats involves processing complex sensory inputs and applying effective decision-making, often requiring an ability to switch strategies77–83. Bats possess a highly accurate spatial memory71,79,83–85, which is essential for tasks like long-distance migration51, homing86, and maneuvering in cluttered environments84. Additionally, they utilize acoustic landmarks to orient in total darkness52, occasionally rely on vision80,81, particularly at the cave edge where light is available, can passively detect echolocating peers, and perhaps eavesdrop on conspecifics’ echoes23. In this study we focused on whether echolocation alone is sufficient for one of the most difficult orientation tasks that bats perform – exiting a roost at high densities without prior knowledge of the roost’s shape, aside from the initial flight direction. Thus, our echolocation-only model, which was based on a five-call integration window during most simulations, probably underestimates real bats’ actual performance which also benefits from additional sensory input and can employ addition navigation strategies by sharing information between each other to coordinate and optimize the routes, such as manifested by swarming intelligence33,87,88.

Revised version: Navigation in bats involves processing complex sensory inputs and applying effective decision-making, often requiring an ability to switch strategies88–94. Bats possess a highly accurate spatial memory82,90,94–96, which is essential for tasks like long-distance migration51, homing97, and maneuvering in cluttered environments95. Additionally, they utilize acoustic landmarks to orient in total darkness52, occasionally rely on vision91,92, particularly at the cave edge where light is available, can passively detect echolocating peers, and perhaps eavesdrop on conspecifics’ echoes23. In this study we focused on whether echolocation alone is sufficient for one of the most difficult orientation tasks that bats perform – exiting a roost at high densities without prior knowledge of the roost’s shape, aside from the initial flight direction. Thus, our echolocation-only model, which was based on a five-call integration window during most simulations, probably underestimates real bats’ actual performance which also benefits from additional sensory input and can employ addition navigation strategies by sharing information between each other to coordinate and optimize the routes, such as manifested by swarming intelligence33,98,99.

Our model highlights the importance of considering sensory interference in animal behavior research and illuminates the impressive capabilities of echolocating bats. Additionally, the model showcases the value of simulations and establishes a framework for future studies on collective movement and swarming animals, and on robotics in complex environments.

Methods

Removed in the revised version: The simulated bats rely solely on echolocation to detect and locate obstacles and other bats by analyzing the sound waves they receive. They emit directional echolocation calls and receive the echoes reflected by roost walls and conspecifics, as well as the calls of conspecifics and the echoes returning from their calls. The bats adjust their flight trajectory and echolocation behavior based on the estimated location of the detected objects (range and angle), which deteriorates upon acoustic interference. The detection of the received signals is based on the mammalian gammatone filter bank receiver, under the assumption that bats can differentiate between the desired detected obstacles, conspecifics’ echoes, and masking signals. We conducted 2D simulations with varying number of bats (from 1 to 100) to analyze the flight trajectories with and without masking interference by conspecifics. In the trials without masking interference the bats successfully detected walls and conspecifics without any hindrance. While real-world bat navigation occurs in 3D space, the 2D framework represents a worst-case scenario for echolocation-based navigation, as it increases effective bat density and limits maneuverability compared to a full 3D environment. This approach provides a conservative test of jamming and collision avoidance while maintaining computational tractability, allowing for extensive simulation runs to explore different variables systematically. For a detailed description of the MATLAB simulation see Mazar & Yovel 20207.

Removed in the revised version: The simulation arena was designed to mimic a roost with a corridor-like layout, measuring 14.5 meters in length and 2.5 meters in width, featuring a right-angle turn located 5.5 meters before the exit (see Figure 1A for a top-down view). All bats started at a random position within a 2 × 1.5 m area at the far end of the cave, each initiating flight within a 0.1-second window in a random direction between –30° and +30° relative to the exit (see Figure 1). They employ a simple navigation algorithm that dynamically adjusts flight direction based on the detected obstacles or conspecifics (Supplementary Figure 1 and Figure 1D). If no obstacles or conspecifics are detected, they continue in a correlated random walk with a maximal turning rate of approximately 30 deg/sec. When obstacles are detected, they are first localized with an error (see below and7). Then, if an opening (i.e., a gap of at least 0.5 m between obstacles) is detected, the bats fly through it, if not, they follow the walls while maintaining a 0.8 m distance from them. When approaching an obstacle too closely (<1.5 m and at an angle <60°), they execute an obstacle avoidance maneuver. Close proximity to another bat (<0.4 m) triggers an avoidance maneuver away from the nearest conspecific. To evaluate the choice of these distances (1.5 m from walls and 0.4 m from other bats), we tested the sensitivity of the model to conspecific avoidance distances ranging from 0.2 to 1.6 meters across bat densities of 2 to 40 bats/3m². We observed only a modest effect on exit probability at the highest density, where exit probability increased slightly from 82% to 88% (p = 0.024, t = 2.25, DF = 958). No significant changes were observed in exit time, collision rate, or jamming probability across other densities or conditions (GLM, with the number of bats and avoidance distance set as fixed explanatory variables, and the outcome variable being one of: exit probability, time-to-exit, collision rate, or jamming probability). These findings confirm that the modeled behavior is largely insensitive to this parameter range.

Removed in the revised version: If the bat collides with a wall, it immediately turns so that its new flight direction is at a 90° angle to the wall. Collisions between conspecifics, which are common in nature and generally not disruptive in low velocities, are not explicitly modeled. Instead, during the collision event the bat keeps decreasing its velocity and changing its flight direction until the distance between bats is above the threshold (0.4 m). We assume that the primary cost of such interactions arises from the effort required to avoid collisions resulting in forced changes in flight’s direction and speed, rather than from the collision itself. Each decision relies on a multi-call integration window that records the estimated locations of detected reflectors from recent echolocation calls (see Supplementary Figure 2A). By default, this window includes the last five calls, and we systematically tested the effect of using between 1 and 10 calls. This algorithm functions without any prior knowledge of the bats’ location or the roost’s structure. To assess performance, we measured the probability of successfully exiting the roost within a 15-second window. The time-based exit limit was chosen because it is approximately twice the average exit time for 40 bats under acoustic interference in our model, allowing bats sufficient time to correct their trajectory and circle back if they missed the exit on the first attempt. This threshold keeps simulation times reasonable while still capturing the key aspects of exit dynamics.

Removed in the revised version: Echolocation behavior and flight speed follow the phases widely reported in insectivorous bats, categorized as “search,” “approach,” and “buzz”89–94 with specific echolocation parameters for Pipistrellus kuhlii (Kuhl’s pipistrelle)59 and Rhinopoma microphyllum (greater mouse-tailed bat)25. The transition distances between these phases were identical for both species (see Table 2) and are based on empirical studies documenting hunting and obstacle avoidance behavior58,89,93–96. In nature, call parameters (Inter Pulse Interval (IPI), call duration, and start and stop frequencies) are primarily shaped by the target distance and echo strength. Accordingly, there is little difference in echolocation between prey capture and obstacles-related maneuvers, aside from intensity adjustments based on target strength 76,96–98. In our study, due to the dense cave environment, the bats are found to operate in the approach phase nearly all of the time, which is consistent with natural cave emergence behavior, where they are navigating through a cluttered environment rather than engaging in open-space search. Our model was designed to remain as simple as possible while relying on conservative assumptions that may underestimate bat performance. If, in reality, bats fine-tune their echolocation calls even earlier or more precisely during navigation than assumed, our model would still conservatively reflect their actual capabilities.

Removed in the revised version: The simulated echolocation call consists of the dominant harmony of the bat’s FM Chirp (1st harmony of the PK and 2nd harmony of the RM). The echolocation signals used in our simulation were modeled as logarithmic FM chirps, implemented using the MATLAB built-in function (e.g., chirp(t, f0, t1, f1, ‘logarithmic’)). This approach aligns with the known nonlinear frequency modulation characteristics of Pipistrellus kuhlii (PK) and Rhinopoma microphyllum (RM). Table 2 provides the specific echolocation parameters used in the model, based on Kalko 199558, and Goldshtein 2025 25. During the search phase, the bats fly at a nominal velocity of 6 m/sec, reducing it by half during the approach phase and continuously adjusting their speed according to the relative direction of the target, using a delayed linear adaptive law7,93,99. The maneuverability of the bats is constrained to a maximum of 4 m/sec², limiting both angular and linear accelerations. Additionally, our model includes random individual variations in terminal frequencies, assuming a normal distribution with a standard deviation of 1 kHz across the bats.

Earlier version: The sound intensity of the echoes generated by the bat’s own calls and those of its conspecifics are calculated using the sonar equation7,100 (pp. 196-198), as shown in Equation 1, geometrical relations are according to Supplementary Figure 5. The received levels of the masking calls are determined by using the Friis transmission equation101, as shown in Equation 2. All signal levels were simulated and reported in dB-SPL, referenced to 0.1 meters from the emitting bat. Bats are modeled acoustically as spherical reflectors with a fixed target strength of –23dB assuming reference distance 1 meter, reflecting sound isotropically. This approximates a sphere with a radius of 0.15 m, corresponding to the approximate wingspan of Rhinopoma microphyllum (RM) 25,102. While target strength can vary with wing posture and body geometry, we chose a representative value within the reported biological range for simplicity and model consistency. Our own measurement of a 3D-printed RM bat yielded a target strength of –32 dB, and a sensitivity analysis (Supplementary Figure 4) showed that model performance was only mildly affected across a wide range of target strengths (see Supplementary Figure 4). This supports the robustness of our approach to different sized bats. Walls are modeled as composites of individual reflectors placed 20 cm apart; each treated as a sphere with a 20 cm radius and a target strength of –22.5dB. For simplicity, in our model, the head is aligned with the body, therefore the direction of the echolocation beam is the same as the direction of the flight. The directivity of the calls and the received echoes is defined by the piston model7,92 with radii of 3 mm for the mouth-gap and 7 mm for the ear. The directivity is not directly influenced by velocity but follows behavioral dependent frequency changes. As the bat transitions from search to approach to buzz phases, it emits higher-frequency calls, leading to increased directivity. This shift coincides with a natural reduction in speed during the approach phase. Echo delays are calculated as the two-way travel time of the signals from the emitter to the target. where,

Revised version: The simulated bats rely solely on echolocation to detect and locate obstacles and other bats by analyzing the sound waves they receive. They emit directional echolocation calls and receive the echoes reflected by roost walls and conspecifics, as well as the calls of conspecifics and the echoes returning from their calls. The bats adjust their flight trajectory and echolocation behavior based on the estimated location of the detected objects (range and angle), which deteriorates upon acoustic interference. The detection of the received signals is based on the mammalian gammatone filter bank receiver, under the assumption that bats can differentiate between the desired detected The sound intensity of the echoes generated by the bat’s own calls and those of its conspecifics are calculated using the sonar equation7,108 (pp. 196-198), as shown in Equation 1, geometrical relations are according to Supplementary Figure 5. The received levels of the masking calls are determined by using the Friis transmission equation109, as shown in Equation 2. All signal levels were simulated and reported in dB-SPL, referenced to 0.1 meters from the emitting bat. Bats are modeled acoustically as spherical reflectors with a fixed target strength of -23dB assuming reference distance 1 meter, reflecting sound isotropically. This approximates a sphere with a radius of 0.15 m, corresponding to the approximate wingspan of Rhinopoma microphyllum (RM) 25,110. While target strength can vary with wing posture and body geometry, we chose a representative value within the reported biological range for simplicity and model consistency. Our own measurement of a 3D-printed RM bat yielded a target strength of –32 dB, and a sensitivity analysis (Supplementary Figure 4) showed that model performance was only mildly affected across a wide range of target strengths (see Supplementary Figure 4). This supports the robustness of our approach to different sized bats. Walls are modeled as composites of individual reflectors placed 20 cm apart; each treated as a sphere with a 20 cm radius and a target strength of -22.5dB. For simplicity, in our model, the head is aligned with the body, therefore the direction of the echolocation beam is the same as the direction of the flight. The directivity of the calls and the received echoes is defined by the piston model7,102 with radii of 3 mm for the mouth-gap and 7 mm for the ear. The directivity is not directly influenced by velocity but follows behavioral dependent frequency changes. As the bat transitions from search to approach to buzz phases, it emits higher-frequency calls, leading to increased directivity. This shift coincides with a natural reduction in speed during the approach phase. Echo delays are calculated as the two-way travel time of the signals from the emitter to the target.

Added in the revised version: Where,

Pr: level of the received signal [SPL]

Pt : level of the transmitted call [SPL]

Pmask : level of the masking signal as received by the bat [SPL]

Earlier version: PecℎoesFromMasking: level of the echoes reflected by conspecifics and received by the bat [SPL]

Revised version: PechoesFromMasking : level of the echoes reflected by conspecifics and received by the bat [SPL]

Gt(ϕ, f): gain of the transmitter (mouth of the bat, piston model), as a function of azimuth and frequency (f) [numeric]

Gr(ϕ, f): gain of the receiver (ears of the bat, piston model) [numeric]

ϕtarget : the angle between the bat and the reflected object [rad]

D: distance between the bat and the target [m]

Earlier version: ϕtxrx, Dtxrx: the angle [rad], and the distance [m] between the transmitting conspecific and the receiving focal bat (from the transmitter’s perspective)

Revised version: : the angle [rad], and the distance [m] between the transmitting conspecific and the receiving focal bat (from the transmitter’s perspective)

Earlier version: ϕrxtx, Drxtx: the angle [rad], and the distance [m] between the receiving bat and the transmitting bat (from the receiver’s perspective)

Revised version: : the angle [rad], and the distance [m] between the receiving bat and the transmitting bat (from the receiver’s perspective)

Earlier version: ϕtx: the angle [rad], between the masking bat and target (from the transmitter’s perspective)

Revised version: : the angle [rad], between the masking bat and target (from the transmitter’s perspective)

αatt(f): atmospheric absorption coefficient for sound [dB/m]

σ(f): SONAR cross-section of the target [m2]

λ: The wavelength of the signal [m]

Earlier version: To maintain model simplicity, we did not incorporate Doppler effects in the echolocation model. While Doppler shifts can affect frequency perception, their impact on jamming and navigation performance is minimal in this context103. Moreover, the inter-individual random signals frequencies were larger than the expected Dopplers. In addition, the model does not assign echoes to earlier calls if their delays exceed the bat’s own Inter-Pulse Interval (IPI), and thus does not simulate pulse-echo ambiguity.

Revised version: To maintain model simplicity, we did not incorporate Doppler effects in the echolocation model. While Doppler shifts can affect frequency perception, their impact on jamming and navigation performance is minimal in this context111. Moreover, the inter-individual random signals frequencies were larger than the expected Dopplers. In addition, the model does not assign echoes to earlier calls if their delays exceed the bat’s own Inter-Pulse Interval (IPI), and thus does not simulate pulse-echo ambiguity.

Earlier version: To model the detection process in the bat’s cochlea, we employed a monoaural filter bank receiver47,104,105 consisting of 80 channels, each with three components: (i) a gammatone filter of order 8, acting as a bandpass filter with center frequencies logarithmically scaled between 10kHz and 80kHz7; (ii) a half-wave rectifier; and (iii) a lowpass filter (Butterworth, fc=8kHz, order=6). Object detection and distance estimation are conducted using Saillant’s method7,47,106, based on the sum of detections in the active channels, see Figure 1C, D. Initially, a de-chirping process calculates the reference frequency-delay by detecting the peak in the response of each channel to the emitted call in a noise-free environment. Subsequently, the received signal, containing both desired echoes and masking sounds, passes through the filter bank. In each channel, all peaks above a threshold level are detected and time-shifted by the de-chirp reference. The detection threshold in each channel was set to the higher of two values: either 7 dB above the noise floor (0 dB-SPL) or the maximum received signal level minus 70 dB, thereby enforcing a dynamic range of 70dB. Peaks from all channels are aggregated in 5 µs windows and convolved with a Gaussian kernel with σ=5 µs. Output peaks that exceed the threshold level, set at 10% of the number of active channels, and fall within a time window of 100µs around the expected delay, are considered successful detections.

Revised version: To model the detection process in the bat’s cochlea, we employed a monoaural filter bank receiver47,112,113 consisting of 80 channels, each with three components: (i) a gammatone filter of order 8, acting as a bandpass filter with center frequencies logarithmically scaled between 10kHz and 80kHz7; (ii) a half-wave rectifier; and (iii) a lowpass filter (Butterworth, fc=8kHz, order=6). Object detection and distance estimation are conducted using Saillant’s method7,47,114, based on the sum of detections in the active channels, see Figure 1C, D. Initially, a de-chirping process calculates the reference frequency-delay by detecting the peak in the response of each channel to the emitted call in a noise-free environment. Subsequently, the received signal, containing both desired echoes and masking sounds, passes through the filter bank. In each channel, all peaks above a threshold level are detected and time-shifted by the de-chirp reference. The detection threshold in each channel was set to the higher of two values: either 7 dB above the noise floor (0 dB-SPL) or the maximum received signal level minus 70 dB, thereby enforcing a dynamic range of 70dB. Peaks from all channels are aggregated in 5 µs windows and convolved with a Gaussian kernel with σ=5 µs. Output peaks that exceed the threshold level, set at 10% of the number of active channels, and fall within a time window of 100µs around the expected delay, are considered successful detections.

Earlier version: To evaluate the impact of acoustic interference, we conducted the detection procedure twice. The first, termed “interference-free detection”, comprised only the desired echoes, with white Gaussian noise at a level of 0 dB-SPL and without masking signals. The second, termed “full detection” comprised the desired echoes, Gaussian noise, and the masking signals. Detected echoes in the full detection were defined by the strongest peak within a four-millisecond window (three milliseconds before and one millisecond after, accounting for forward and backward masking 24,107–109) detected above the threshold within 100µs of the interference-free detections. If the detected peak in the full detection condition was delayed by more than 100 µs compared to the interference-free case, it was defined as a miss-detection. Peaks with smaller timing shifts were considered detection with timing errors. Jammed echoes were defined as echoes that were detected under the interference-free condition but not detected under the full detection condition. The jamming probability was calculated as the ratio of jammed echoes in the full detection condition to the detected echoes in the interference-free condition.

Revised version: To evaluate the impact of acoustic interference, we conducted the detection procedure twice. The first, termed “interference-free detection”, comprised only the desired echoes, with white Gaussian noise at a level of 0 dB-SPL and without masking signals. The second, termed “full detection” comprised the desired echoes, Gaussian noise, and the masking signals. Detected echoes in the full detection were defined by the strongest peak within a four-millisecond window (three milliseconds before and one millisecond after, accounting for forward and backward masking 24,115–117) detected above the threshold within 100µs of the interference-free detections. If the detected peak in the full detection condition was delayed by more than 100 µs compared to the interference-free case, it was defined as a miss-detection. Peaks with smaller timing shifts were considered detection with timing errors. Jammed echoes were defined as echoes that were detected under the interference-free condition but not detected under the full detection condition. The jamming probability was calculated as the ratio of jammed echoes in the full detection condition to the detected echoes in the interference-free condition.

Earlier version: After detection, the bat estimates the range and the Direction of Arrival (DOA) of the reflecting objects. The range is determined by the delay of the detected echo, including any errors derived from the filter-bank process in the “full detection” process (i.e., including all masking signals).7,110,113. The direction is not explicitly estimated through binaural processing. Instead, based on previous studies 107,110, we assumed that bats can estimate the direction of arrival with an angular error that depends on the Signal-to-Noise Ratio (SNR) and the angle. The inputs to this process include the peak level of the desired echo, the noise level, and the level of acoustic interference. The output is the estimated direction of arrival with a random error applied based on the SNR. At an angle of 0° and an SNR of 10 dB, the standard deviation of the error is 1.5° 111 and 7 (Equation 10).

Revised version: After detection, the bat estimates the range and the Direction of Arrival (DOA) of the reflecting objects. The range is determined by the delay of the detected echo, including any errors derived from the filter-bank process in the “full detection” process (i.e., including all masking signals).7,110,113. The direction is not explicitly estimated through binaural processing. Instead, based on previous studies 115,118, we assumed that bats can estimate the direction of arrival with an angular error that depends on the Signal-to-Noise Ratio (SNR) and the angle. The inputs to this process include the peak level of the desired echo, the noise level, and the level of acoustic interference. The output is the estimated direction of arrival with a random error applied based on the SNR. At an angle of 0° and an SNR of 10 dB, the standard deviation of the error is 1.5° 119 and 7 (Equation 4), with the error capped at a maximum of 3° in our model. where, k₂, k₃, and k₄ are constants chosen to produce a DOA error consistent with the range described above.

To evaluate the impact of the assumption that bats can distinguish between echoes caused by their own calls and those caused by other bats (i.e., conspecifics’ reflectors), we tested an alternative model in which the simulated bats treat all echoes reflected from walls as if they have originated from their own calls The distance to reflectors of conspecifics’ calls is estimated based on the time difference between the echo and the bat’s last call. The direction of arrival is estimated by the angle between the bat and the physical reflector, with an added random error (the same process used for their own echoes).

Earlier version: In real bats, spatial processing in the brain involves integrating auditory and spatial information over time to construct a coherent map of the environment 5,57. This neural computation is crucial for navigation and prey detection in complex environments. To examine whether spatial integration mitigates the confusion problem, we added a ‘multi-call clustering’ module that was based on the sensory information obtained within a one-second memory window. The clustering comprised the following steps: (i) clustering all detections in memory into groups with a maximum internal distance of 10 cm; (ii) reconstructing the estimated walls positions and directions based on the average of clusters that include at least two detections (rather than relying on single reflections); and (iii) identifying openings between reconstructed wall edges ranging from 0.5 to 2.25 meters in width, see Supplementary Figure 1 and Supplementary Figure 2B. The model assumes that bats store echo locations in an allocentric x-y coordinate system, transforming detections from a local to a global spatial framework. Collision avoidance is based not only on the integrated spatial representation but also on immediate echoes from the last call (prior to clustering), including potential uncorrected false detections and localization errors, which are independently processed for real-time evasive maneuvers.

Revised version: In real bats, spatial processing in the brain involves integrating auditory and spatial information over time to construct a coherent map of the environment 5,68. This neural computation is crucial for navigation and prey detection in complex environments. To examine whether spatial integration mitigates the confusion problem, we added a ‘multi-call clustering’ module that was based on the sensory information obtained within a one-second memory window. The clustering comprised the following steps: (i) clustering all detections in memory into groups with a maximum internal distance of 10 cm; (ii) reconstructing the estimated walls positions and directions based on the average of clusters that include at least two detections (rather than relying on single reflections); and (iii) identifying openings between reconstructed wall edges ranging from 0.5 to 2.25 meters in width, see Supplementary Figure 1 and Supplementary Figure 2B. The model assumes that bats store echo locations in an allocentric x-y coordinate system, transforming detections from a local to a global spatial framework. Collision avoidance is based not only on the integrated spatial representation but also on immediate echoes from the last call (prior to clustering), including potential uncorrected false detections and localization errors, which are independently processed for real-time evasive maneuvers.

Statistical analysis

Statistical analysis and the roost-exit model were conducted using MATLAB© 2023a.

Tests were performed with a significance level of 0.05. For each simulated scenario, we examined the effect of the various parameters on exit probability, time-to-exit, collision rate, and the jamming probability, using Generalized Linear Models (GLMs). The GLM tests were executed with MATLAB built-in function ‘fitglm()’. Probability variables (such as exit and jamming probabilities) were treated as binomially distributed; rate variables (such as collision rate) were treated as Poisson distributed, and all other variables were considered normally distributed. Unless otherwise stated, all explaining variables were set as fixed factors. All statistical analyses, including the statistical test and the corresponding sample sizes, are described throughout the text and summarized in Table 1. Standard errors are calculated across all individuals within each scenario, without distinguishing between different simulation trials.

Removed in the revised version: Data availability

Removed in the revised version: All data and codes generated during this study are included in the manuscript and supporting files. Source code files have been uploaded with a Graphical User Interface and a readme file for explanation. Data are available at: https://github.com/omermazar/Colony-Exit-Bat-Simulation/tree/main

Removed in the revised version: Supplementary material

Removed in the revised version: Decision-making in echolocation-based pathfinding

Removed in the revised version: This diagram illustrates the sensorimotor decision-making process based solely on echolocation. The process starts with the emission of an echolocation call (1) and the reception of echoes and interfering signals (2), followed by sensory processing for detection, range estimation, and direction of arrival (DOA) (3). After integrating detections over a 1–10 call window (4), the bat engages in crash avoidance (5) by evaluating the proximity of conspecifics and obstacles directly ahead. If either is too close, the bat turns in the opposite direction of the detected obstacle, by applying maximum angular velocity away from it (e.g., if the obstacle is on the right, the bat turns left). If no immediate threat is detected, the bat proceeds to pathfinding (6). During pathfinding, it checks for obstacles and, if an opening is detected, flies toward the gap’s center. Without the optional multi-call clustering process (green), the bat simply integrates detections and flies toward the farthest detected obstacle, interpreting it as a wall edge. If the multi-call clustering is included (9), a one-second short memory aids in clustering detections, estimating wall edges, and identifying openings, while also allowing the bat to follow walls at a constant distance. Throughout, the bat continuously adjusts echolocation parameters (8) and controls flight direction and velocity (7) based on ongoing sensory information and decision-making.

Removed in the revised version: Multi-Call Integration

Removed in the revised version: This figure demonstrates the effect of multi-call integration under non-confusing conditions. The upper-right panel shows the position of the focal bat (black) and nine conspecifics (red) within the roost corridor, with a zoomed-in view of the gray rectangle provided in Panels A–C. (A) When the integration window is set to zero calls (no memory), the bat relies solely on the latest call. Green circles and squares represent detected reflectors, while red Xs indicate missed (jammed) detections. Notably, the left wall of the corridor remains undetected due to jamming. (B, C) Increasing the integration window to five calls (magenta squares) and ten calls (black squares) accumulates detections from prior calls, improving coverage of the environment. In this basic integration model, each detection is treated independently, without clustering. (D1, D2) Magnified views of the grey regions indicated in Panel C, comparing detections across 0, 5, and 10-call windows (green, magenta, and black, respectively), illustrating how extended memory improves detection robustness. Note that the X-Y aspect ratios in D1 and D2 differ from the main panels to enhance visibility of spatial distributions.

Removed in the revised version: Multi-Call Clustering Example

Removed in the revised version: This figure illustrates the multi-call clustering algorithm under full-confusion conditions. (A) A focal bat (black) and four conspecifics (red) are shown in the lower corridor. (B) A zoom-in of the gray rectangle in (A). Black ovals represent detections from the last call; red X’s indicate jammed echoes; black squares represent all detections stored across the integration window (before clustering), each subject to localization error. When not applying multi-call clustering – the bat would rely on all of these dots as reflectors. Under full confusion, the bat cannot distinguish self-echoes from conspecific echoes, leading to false detections (red diamonds). Detections are clustered when a reflector is detected twice or more within a 10 cm radius (green squares). The clustered reflectors are used to estimate wall directions (blue dashed line) and detect possible gaps (not shown). As a result of to the multi-call clustering algorithm, most false detections are removed as outliers, except for one erroneous cluster (Panel A). Collision avoidance maneuvers are based on both the clustered obstacles and the raw detections from the latest call (empty black ovals).

Removed in the revised version: Sensitivity of exit performance to obstacle target strength

Removed in the revised version: This figure shows how changes in the acoustic target strength of the cave walls affect navigation performance across five bat densities (1, 5, 10, 40, and 100 bats/3 m²). Target strength values ranged from –33 dB to –3 dB, corresponding to spherical reflectors with approximate radii from 0.05 m to 1.5 m. Overall, increasing obstacle target strength significantly influenced exit performance, primarily by reducing the probability of obstacle jamming and thereby improving detection. (A) Exit Probability increased with obstacle target strength across all densities, with a maximal increase of 64% for a density of 100 bats (p << 10⁻¹⁰, t = 28.5, DF = 8157, GLM). (B) Time to Exit decreased significantlywith increasing obstacle target strength, with a maximal reduction of approximately 2.8 seconds at a density of 10 bats (p << 10⁻¹⁰, t = –22.2, DF = 6920, GLM). (C) Conspecific Collision Rate increased slightly with stronger obstacle echoes (p << 10⁻¹⁰, t = 27.6, DF = 8157, GLM). (D) Obstacle Collision Rate decreased significantly with increasing target strength (p << 10⁻¹⁰, t = –10.7, DF = 8157, GLM), reflecting better detection of walls and structures. (E) Obstacle Jamming Probability decreased consistently (p << 10⁻¹⁰, t = –19.8, DF = 8157, GLM). (F) Conspecific Jamming Probability increased with obstacle target strength (p << 10⁻¹⁰, t = 27.6, DF = 8157, GLM). These results suggest that stronger wall echoes improve environmental awareness at the cost of slightly increased masking of conspecific echoes. Despite this, the overall performance— particularly exit probability and reduced obstacle collisions—improves significantly. In all panels, circles represent means and bars represent standard errors. The error bars are present but very small due to the large number of simulation repetitions, and thus may not be visually noticeable at the plotted scale. See Table 1 for the number of simulated bats.

Removed in the revised version: Sensitivity of exit performance to conspecific’s target strength

Removed in the revised version: This figure shows how changes in the acoustic target strength of conspecifics affect navigation performance across four bat densities (1, 5, 10, and 40 bats/3 m²). Overall, our results indicate that target strength has a relatively minor impact on performance, likely because it affects both desired echo signals and masking signals equally. Interestingly, this analysis also suggests that our model is more sensitive to the bat’s response to nearby conspecifics than to the physical collision impact itself. (A) Exit probability was not significantly affected by conspecific target strength (p=0.28, t=-1.09, DF=5757, GLM, see details in Table 1). Note that the performance curves for densities of 1 and 5 bats overlap almost completely. (B) Time-to-exit increased with target strength at high density, with a maximal effect size of ∼1 second at 40 bats (p = 0.003, t = 3.02, DF = 5578). (C, D) Collision rates with conspecifics decreased significantly with stronger target strength (p = 0.0002, t = –3.7, DF = 5757), while collisions with obstacles remained statistically unchanged (p = 0.23, t = 1.18, DF = 5757). (E, F) Jamming probability was not significantly affected for either conspecific or obstacle echoes (p = 0.6, t = –0.51, DF = 4762; p = 0.19, t = 1.31, DF = 5757, respectively). This aligns with the notion that both useful and interfering signals scale similarly with target strength. Importantly, the probability of detecting a conspecific located within 1 meter increased substantially with higher target strength, improving from 25% to 43% at 40 bats (p < 10⁻¹⁰, t = 6.45, DF = 4162). In all panels, circles represent means and bars represent standard errors. The error bars are present but very small due to the large number of simulation repetitions, and thus may not be visually noticeable at the plotted scale. See Table 1 for the number of simulated bats.

Removed in the revised version: Angles and distances for two bats and two reflecting objects.

Removed in the revised version: Bat1 receives a reflected echo from Prey1 or a stationary obstacle located at a distance of D from it, with an angle ϕtarget relative to its flight direction (red arrow 1). Prey1 is also within the detection range of Bat1, depicted by the green shaded piston area. Bat1 also receives masking sounds from Bat2. The echolocation signals emitted by Bat2 arrive at the ear of Bat1 at an angle ϕtxrx relative to its flight direction and from a distance of Dtxrx (red arrow 2). Additionally, the echolocation signals of Bat2 are reflected by Prey2, before being received by Bat 1. These reflected signals act as masking signals at a relative angle of angle ϕrx, and from a distance of Drx from Bat1.